The autophagy protein, ATG14 safeguards against unscheduled pyroptosis activation to enable embryo transport during early pregnancy

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

This study by Popli et al. evaluated the function of Atg14, an autophagy protein, in reproductive function using a conditional knockout mouse model. The authors showed that female mice lacking Atg14 were infertile partly due to defective embryo transport function of the oviduct and faulty uterine receptivity and decidualization using PgrCre/+;Atg14f/f mice. The findings from this work are exciting and novel. The authors demonstrated that a loss of Atg14 led to an excessive pyroptosis in the oviductal epithelial cells that compromises cellular integrity and structure, impeding the transport function of the oviduct. In addition, the authors use both genetic and pharmacological approaches to test the hypothesis. Therefore, the findings from this study are high-impact and likely reproducible. However, there are multiple major concerns that need to be addressed to improve the quality of the work.

Thank you for the additional data that solidified the conclusion of this study. The authors addressed almost all of my previous concerns in this revised manuscript. However, some key points wording still need to be addressed.

Comments on revisions:

In Fig. 2A, please ensure that these are 5.0 dpc samples since implantation has already occurred at this point. However, the embryo appeared free-floating adjacent to the luminal epithelial cells (LE), even in control.

We understand the reviewer’s concern. We have now replaced the previous H & E image with a clearer, higher-quality section that shows a fully attached embryo within a closed uterine lumen representing a typical implantation morphology at the D5 stage of pregnancy. (Revised Figure 2A)

Fig. 3A-B: "Approximately 80-90% of blastocysts" contradicts the quantification in Figure 3C, which showed a percentage of blastocysts below 50%. Please clarify and correct as needed.

In Fig. 3A-B, we mean to say approximately 80-90% embryos. We have now corrected the statement in the revised manuscript (Line no: 349-351).

The authors showed that Acetylated a-tubulin was present in the ampulla region of cKO (Fig. 4A). However, the revised manuscript still stated that (lines 397-399) ...there was a substantial loss of the ciliary epithelial cells (indicated by fewer a-tubulin and FOXJ1-positive cells) (Fig. 4B, left panel and Fig. S3)... So, the authors may want to tone down their conclusion regarding a "substantial loss" of ciliated epithelial cells if the quantification of ciliated cell number is not performed.

We thank the reviewer for this suggestion. To avoid redundancy and ambiguity, we have revised the statement as below (Line no: 391-395):

“As shown in Fig. 4A, normal ciliary structures were observed in the ampulla of both control and cKO oviducts. However, in the isthmus of cKO oviducts, we observed a reduction in both the FOXJ1- and PAX8-expressing cells (Fig. 4B, and Fig. S3).”

Fig. 4C - the areas with red inset boxes labeled for isthmus are not really isthmus (in both control and cKO). The zoomed-in images (Fig. 4C - The far-right panel for both control and cKO, images are the transitional zone from the ampulla to the isthmus. The isthmus areas should have a thick muscle layer with almost no ciliated cells - see Fig. 4B cKO - those are true isthmus areas.

We thank the reviewer for noting this. We have corrected the label accordingly. Since ciliary epithelial cells predominantly reside in the ampulla, we have included high-resolution images specifically for the ampulla regions.

• Fig. 3A and 3C, it appears that the images were taken at different magnifications, but the scale bars are the same at 200 um. The authors, please double-check the scale bars.

We thank the reviewer for noting this. We have double-checked all the figures to ensure the scale bars are correctly displayed and aligned with the resolution.

• Fig. 6D - why polyphillin-treated samples did not sum to 100%? - please double-check.

Since approximately 50% of the embryos were retained in the oviduct following polyphyllin treatment (Figure 6C, upper bar), the bar in Figure 6D represents this percentage (50% retained) rather than 100%.

Reviewer #2 (Public review)

In this manuscript, Popli et al investigated the roles of autophagy-related gene, Atg14, in the female reproductive tract (FRT) using conditional knockout mouse models. By ablation of Atg14 in both oviduct and uterus with PR-Cre (Atg14 cKO), authors discovered that such females are completely infertile. They went on to show that Atg14 cKO females have impaired embryo implantation as well as embryo transport from oviduct to uterus. Further analysis showed that Atg14 cKO leads to increased pyroptosis in oviduct, which disrupts oviduct epithelial integrity and leads to obstructive oviduct lumen and impaired embryo transport. The authors concluded that Atg14 is critical for maintaining the oviduct homeostasis and keeping the inflammation under check to enable proper embryo transport.

The authors have barely addressed most of my concerns in this revised version with a few minor issues remaining to be addressed:

(1) The authors tried to address my first concern regarding the statement that "autophagy is critical for maintaining the oviduct homeostasis". The revised statement in Lines 53-54 "we report that Atg14-dependent autophagy plays a crucial role in maintaining..." is still not correct. It should be corrected as " we report that autophagy-related protein Atg14 plays a crucial role in maintaining...".

We thank the reviewer for this nice suggestion. We have now modified the statement as suggested (Line no: 54).

(2) Line 349-351 described 80-90% of blastocysts retrieved from oviducts of cKO mice, which is in consistent with Figure 3B (showing more than 98%).

We thank the reviewer for noting this. We have now corrected the statement as: “Unexpectedly, oviduct flushing from cKO mice resulted in the retrieval of approximately 90% of embryos, suggesting their potential entrapment within the oviducts, impeding their transit to the uterus”. (Line No: 349-351).

(3) Line 447, "Fig. 5E" should be Fig. 6A. In addition, grammar error in the next sentence.

We have corrected the figure number and addressed the grammatical error.

(4) In Figure 6D, why the composition of blastocysts in chemical treated group do not add up to 100%.

As explained in Reviewer 1 responses, the bar in Figure 6D represents the 50% retained embryos from Figure 6C upper bar the full count.

Reviewer #3 (Public review):

Summary:

The manuscript by Pooja Popli and co-authors tested the importance of Atg14 in the female reproductive tract by conditionally deleting Atg14 use PrCre and also Foxj1cre. The authors showed that loss of Atg14 leads to infertility due to the retention of embryos within the oviduct. The authors further concluded that the retention of embryos within the oviduct is due to pyroptosis in oviduct cells leading to defective cellular integrity. The revised manuscript has included new experimental data (Figs. S2B, 5B, 5C, and S3) that satisfied the concerns of this reviewer. The manuscript should provide important advancement to the field.

We sincerely thank the reviewer for the thoughtful evaluation of our manuscript and appreciate your constructive feedback.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Recurrent pregnancy loss (RPL), characterized by two or more failed clinical pregnancies, poses a significant challenge to reproductive health. In addition to embryo quality and endometrial function, proper oviduct function is also essential for successful pregnancy establishment. Therefore, structural abnormalities or inflammation resulting from infection in the oviduct may impede the transport of embryos to the endometrium, thereby increasing the risk of miscarriage. However, our understanding of the biological processes that preserve the oviductal cellular structure and functional integrity is limited. Here, we report that Atg14-dependent autophagy plays a crucial role in maintaining the cellular integrity of the oviduct by controlling inflammatory responses, thereby supporting efficient embryo transport. Specifically, the conditional depletion of the autophagy-related gene, Atg14 in the oviduct causes severe structural abnormalities compromising its cellular integrity leading to the abnormal retention of embryos. Interestingly, the selective loss of Atg14 in oviduct ciliary epithelial cells did not impact female fertility, highlighting the specificity of ATG14 function in distinct cell types within the oviduct. Mechanistically, loss of Atg14 triggered unscheduled pyroptosis via altering the mitochondrial integrity leading to inappropriate embryo retention and impeded embryo transport in the oviduct. Finally, pharmacological activation of pyroptosis in pregnant mice phenocopied the genetically induced defect and caused impairment in embryo transport. Together, we found that ATG14 safeguards against unscheduled pyroptosis activation to enable embryo transport from the oviduct to uterus for the successful implantation. Of clinical significance, these findings provide possible insights into the underlying mechanism(s) of early pregnancy loss and might aid in developing novel prevention strategies using autophagy modulators.

Revised version: Recurrent pregnancy loss (RPL), characterized by two or more failed clinical pregnancies, poses a significant challenge to reproductive health. In addition to embryo quality and endometrial function, proper oviduct function is also essential for successful pregnancy establishment. Therefore, structural abnormalities or inflammation resulting from infection in the oviduct may impede the transport of embryos to the endometrium, thereby increasing the risk of miscarriage. However, our understanding of the biological processes that preserve the oviductal cellular structure and functional integrity is limited. Here, we report that autophagy-related protein ATG14 plays a crucial role in maintaining the cellular integrity of the oviduct by controlling inflammatory responses, thereby supporting efficient embryo transport. Specifically, the conditional depletion of the autophagy-related gene, Atg14 in the oviduct causes severe structural abnormalities compromising its cellular integrity leading to the abnormal retention of embryos. Interestingly, the selective loss of Atg14 in oviduct ciliary epithelial cells did not impact female fertility, highlighting the specificity of ATG14 function in distinct cell types within the oviduct. Mechanistically, loss of Atg14 triggered unscheduled pyroptosis via altering the mitochondrial integrity leading to inappropriate embryo retention and impeded embryo transport in the oviduct. Finally, pharmacological activation of pyroptosis in pregnant mice phenocopied the genetically induced defect and caused impairment in embryo transport. Together, we found that ATG14 safeguards against unscheduled pyroptosis activation to enable embryo transport from the oviduct to uterus for the successful implantation. Of clinical significance, these findings provide possible insights into the underlying mechanism(s) of early pregnancy loss and might aid in developing novel prevention strategies using autophagy modulators.

Introduction

A successful pregnancy is orchestrated by the sequential and coordinated events happening in the FRT [1]. Each of these events is crucial to advance to the next step in pregnancy. For example, as a first step, the ovary undergoes ovulation to release a mature ovum which is collected by oviduct fimbriae and funnelled through the infundibulum into the ampulla where they get fertilized by sperm. The embryos then pass through the ampulla-isthmus junction (AIJ) to enter the isthmus of the oviduct before exiting via the distal utero-tubal junction (UTJ) to implant on the wall of the uterus. Recently, the intricate structural and cellular diversity of the oviduct has drawn significant attention, emphasizing its vital role in facilitating embryo development and transport [2–4]. Morphologically, the oviduct is divided into four evolutionarily conserved regions: the infundibulum (nearest the ovary), ampulla (the site of fertilization), isthmus (serving as a sperm reservoir and site for early embryonic development), and utero-tubal junction (connected to the uterus). Each of these segments comprises epithelial, stromal, and smooth muscle cells with lumen lined up with an epithelium containing secretory (PAX8+) and multi-ciliated (FOXJ1+) cells. While the infundibulum and ampulla are composed of more ciliated than secretory cells, the isthmus has more secretory than ciliated cells [5, 6]. Successful pregnancy requires that the embryo transit the entire oviduct orchestrated by the following: (1) secretory cells producing oviduct fluid, (2) beating cilia to ensure unidirectional flow of the fluid, and (3) periodic contractions of surrounding muscle to propel fluid flow. Therefore, a combination of ciliary and muscular activity contributes to the overall success of oviduct embryo transport. This combination of ciliary and muscular activity plays a crucial role in the effective transport of the embryo through the oviduct. Any perturbations in these early pregnancy events can lead to adverse ripple effects that can compromise pregnancy outcomes. Previous studies have reported that impaired oviductal transport of embryos can lead to pregnancy failure and cause infertility in mice [7–10]. However, the underlying mechanism is not yet completely understood.

Autophagy is a cellular process, evolutionarily conserved from yeast to mammals, that recycles long-lived proteins and organelles to maintain cell energy homeostasis. Autophagy is classically activated through the nutrient sensor or mammalian target of the rapamycin complex. Genetic screens for autophagy-defective mutants in yeast and other fungi have currently identified 41 autophagy-related (ATG) genes that play a primary role in autophagy [11, 12]. Approximately half of these genes have homologs in higher organisms, and Atg14 (also known as Barkor for Beclin 1 (Becn1)-associated autophagy-related key regulator) is one of them [13, 14]. ATG14 is part of a protein complex that is composed of Beclin 1, vacuolar sorting protein 15 (VPS15), and VPS34 (also named as Pik3c3, the catalytic subunit of the class III phosphatidylinositol 3-kinase), and this ATG14-containing complex plays an important role in the initiation process of autophagy.

Pyroptosis is a type of inflammatory cell death mediated by gasdermin (GSDM) and is a product of continuous cell expansion until the cytomembrane ruptures, resulting in the release of cellular contents that can activate strong inflammatory and immune responses. Gasdermin family proteins are the primary executioners of pyroptosis. Cytotoxic N-terminal of gasdermins generated from caspases mediated cleavage of gasdermin proteins oligomerizes and forms pores across the cell membrane, leading to the release of proinflammatory cytokines such as interleukin-1β (IL-1β) and interleukin-18 (IL-18). Aberrant activation of pyroptosis has been implicated in the progression of many diseases including cancer and autoimmune, cardiovascular, and infectious diseases [15–17]. Recent studies have shown that the interactions between pyroptosis and autophagy play an important role in various physiological and pathological processes [15–20]. For example, autophagy inhibition upregulates galangin-induced pyroptosis in human glioblastoma multiforme cells and promotes pneumococcus-induced pyroptosis [21, 22]. However, the studies defining the molecular interactions between the pyroptosis, and autophagy pathways are very limited.

Earlier version: With a better understanding of the fundamental process of autophagy, its pathophysiological functions have begun to be appreciated in the female reproductive tract. For example, mice deficient in key autophagy genes such as genetic knockout of Atg7 or Becn1 result in primary ovarian insufficiency and reduced progesterone production [23, 24]. Similarly, recent studies from our group established the roles of three different autophagy-specific genes: ATG16L, FIP200/RB1CC1, or BECN1 in endometrial physiological processes, including receptivity and decidualization [25–27]. However, none of these proteins exhibited any discernible impact on oviduct function. Surprisingly, in this study, we revealed a critical role for Atg14 in maintaining proper oviduct function, specifically enabling the transport of embryos to the uterus-a function distinct from that of other autophagy-related proteins. Loss of ATG14 in the oviduct resulted in severe structural abnormalities, compromising its cellular integrity, ultimately leading to embryo retention and infertility.

Revised version: With a better understanding of the fundamental process of autophagy, its pathophysiological functions have begun to be appreciated in the female reproductive tract. For example, mice deficient in key autophagy genes such as genetic knockout of Atg7 or Becn1 result in primary ovarian insufficiency and reduced progesterone production [23, 24]. Similarly, recent studies from our group established the roles of three different autophagy-specific genes: Atg16l, Fip200/Rb1cc1, or Becn1 in endometrial physiological processes, including receptivity and decidualization [25–27]. However, none of these genes exhibited any discernible impact on oviduct function. Surprisingly, in this study, we revealed a critical role for ATG14 in maintaining proper oviduct function, specifically enabling the transport of embryos to the uterus- a function distinct from that of other autophagy-related proteins. Loss of ATG14 in the oviduct resulted in severe structural abnormalities, compromising its cellular integrity, and ultimately leading to embryo retention and infertility.

Materials and methods

Animal care and use

Earlier version: All animal studies were approved by the Institutional Animal Care and Use Committee of Washington University School of Medicine, Saint Louis, MO, USA and Use Committee of Baylor College of Medicine, Houston, TX, USA. Atg14 f/f mice were provided by gift from Dr. Shizuo Akira at the Department of Host Defense, Research Institute for Microbial Diseases (RIMD), Osaka University, and previously described [28]. Wild type Pr-cre mice were provided by Dr. John Lydon at Baylor College of Medicine, Houston and previously described [29]. Atg14flox/flox mice, in which exon 4 was flanked by loxp sites, were bred to progesterone receptor cre (PRcre/+) mice to generate (Atg14flox/flox; PRcre/+ mice), hereafter referred to as Atg14 cKO mice. Both control and conditional knockout females were generated by crossing females carrying homozygous Atg14flox/flox alleles with Atg14 cKO males. Foxj1-cre mice were a generous gift from Dr. Michael J. Holtzman at Washington University St. Louis and previously described [30]. Foxj1/Atg14 cKO mice were generated by crossing females carrying homozygous Atg14 f/f/Foxj1f/f females with Atg14f/f/Foxj1-cre males. All mice were age-matched and on a C57BL/6 genetic background (The Jackson Laboratory, Bar Harbor, ME). Mice were genotyped by PCR analysis of genomic DNA isolated from tail clippings using the gene-specific primers listed in Supplementary Table 1.

Revised version: All animal studies were approved by the Institutional Animal Care and Use Committee of Washington University School of Medicine, Saint Louis, MO, USA, and the Use Committee of Baylor College of Medicine, Houston, TX, USA. Atg14 f/f mice were provided by a gift from Dr. Shizuo Akira at the Department of Host Defense, Research Institute for Microbial Diseases (RIMD), Osaka University, and previously described [28]. Wild-type Pr-cre mice were provided by Dr. John Lydon at Baylor College of Medicine, Houston, and previously described [29]. Atg14flox/flox mice, in which exon 4 was flanked by loxp sites, were bred to progesterone receptor cre (PRcre/+) mice to generate (Atg14flox/flox; PRcre/+ mice), hereafter referred to as Atg14 cKO mice. Both control and conditional knockout females were generated by crossing females carrying homozygous Atg14flox/flox alleles with Atg14 cKO males. Foxj1-cre mice were a generous gift from Dr. Michael J. Holtzman at Washington University St. Louis and previously described [30]. Foxj1/Atg14 cKO mice were generated by crossing females carrying homozygous Atg14 f/f/Foxj1f/f females with Atg14f/f/Foxj1-cre males. All mice were age-matched and on a C57BL/6 genetic background (The Jackson Laboratory, Bar Harbor, ME). Mice were genotyped by PCR analysis of genomic DNA isolated from tail clippings using the gene-specific primers listed in Supplementary Table 1.

Fertility analysis and timed mating

Female fertility was determined by mating cohorts of Atg14 cKO experimental (n=6) and control Atg14 f/f (n=4) females individually starting at 8 weeks of age with sexually mature males of proven fertility. Similarly, breeding trials were set up for Foxj1/Atg14-cre females (n=6) and Foxj1/Atg14 control 8-week-old f/f females (n=6). The numbers of litters and pups were tracked over 6 months for each female. Pups per litter for each genotype are reported as mean ± SEM. For timed mating, the morning on which the copulatory plug was first observed was considered 1 dpc. To visualize implantation sites, mice received a tail vein injection of 50 μL of 1% Chicago Sky Blue dye (Sigma-Aldrich, St. Louis, MO, USA) at 5 dpc just before sacrifice.

Steroid hormone treatments

The hormonal profile of pregnancy at the time of implantation was done using a previously described experimental scheme [26]. Briefly, Atg14 cKO and control females (8 weeks old) were bilaterally ovariectomized under ketamine anesthesia with buprenorphine-SR as an analgesic. Mice were allowed to rest for two weeks to dissipate all endogenous ovarian hormones. After the resting period, mice were injected with 100 ng of estrogen (E2; Sigma-Aldrich) dissolved in 100 µL of sesame oil on two consecutive days and then allowed to rest for two days. At this point, mice were randomly divided into three groups of five: Vehicle-treated (E2 priming) mice received four consecutive days of sesame oil injections; E2 group mice received three days of sesame oil injections followed by a single injection of 50 ng of E2 on the fourth day; The E2/P4 mice received 1 mg of progesterone (P4; Sigma-Aldrich) for three consecutive days followed by a single injection of 1 mg P4 plus 50 ng E2 on the fourth day. All hormones were delivered by subcutaneous injection in a 90:10 ratio of sesame oil: ethanol. Mice were euthanized 16 hours after the final hormone injection to collect the uteri. A small piece of tissue from one uterine horn was processed in 4% neutral buffered paraformaldehyde for histology, and the remaining tissue was snap-frozen and stored at −80°C.

Hormone Analysis

Earlier version: For serum hormone levels measurement, blood was collected from D4 pregnant mice before mice were sacrificed. Serum was separated from the blood by centrifugation and stored at −80°C before hormone analysis. Serum P4 and E2 levels were measured by using ELISA kits (Enzo life Sciences) according to the manufacturer’s instructions.

Revised version: For serum hormone level measurement, blood was collected from D4 pregnant mice before mice were sacrificed. Serum was separated from the blood by centrifugation and stored at - 80°C before hormone analysis. Serum P4 and E2 levels were measured by using ELISA kits (Enzo Life Sciences) according to the manufacturer’s instructions.

Hematoxylin and eosin staining

Tissues were fixed in 4% paraformaldehyde, embedded in paraffin, and then sectioned (5 μm) with a microtome (Leica Biosystem, Wetzlar, Germany). Tissue sections were deparaffinized, rehydrated, and stained with Hematoxylin and Eosin (H&E) as described previously [31]. All the histology was performed on three sections from each tissue of individual mice, and one representative section image is shown in the respective figures.

Histological analysis

For histological analysis, the collected tissues (oviduct or uteri) were fixed in 4% paraformaldehyde, embedded in paraffin, and sectioned. Sections (5 µm) were immunostained (n=5 per group) as described previously [31]. Briefly, after deparaffinization, sections were rehydrated in an ethanol gradient and then boiled for 20 min. in citrate buffer (Vector Laboratories Inc., Newark, CA, USA) for antigen retrieval. Endogenous peroxidase activity was quenched with Bloxall (Vector Laboratories Inc.), and tissues were blocked with 2.5% goat serum in PBS for 1 hr (Vector Laboratories Inc.). After washing in PBS three times, tissue sections were incubated overnight at 4 °C in 2.5% goat serum containing the primary antibodies listed in Supplementary Table 2. Sections were incubated for 1 hr with biotinylated secondary antibody, washed, and incubated for 45 min with ABC reagent (Vector Laboratories Inc.). Color was developed with 3, 3’-diaminobenzidine (DAB) peroxidase substrate (Vector Laboratories Inc.), and sections were counter-stained with hematoxylin. Finally, sections were dehydrated and mounted in Permount histological mounting medium (Thermo Fisher Scientific, Waltham, MA, USA).

Transmission electron microscopy (TEM)

For ultrastructural analysis, oviducts were fixed in 2% paraformaldehyde/2.5% glutaraldehyde (Ted Pella Inc., Redding, CA) in 100 mM cacodylate buffer, pH 7.2 for 1 hr at room temperature and then overnight at 4°C. Samples were washed in cacodylate buffer and postfixed in 1% osmium tetroxide (Ted Pella Inc.) for 1 hr. Samples were then rinsed extensively in dH20 prior to en bloc staining with 1% aqueous uranyl acetate (Ted Pella Inc.) for 1 hr. Following several rinses in dH20, samples were dehydrated in a graded series of ethanol and embedded in Eponate 12 resin (Ted Pella Inc.). For initial evaluation semithin sections (0.5 μm) were cut with a Leica Ultracut UCT7 ultramicrotome (Leica Microsystems Inc., Bannockburn, IL) and stained with toluidine blue. Sections of 95 nm were then cut and stained with uranyl acetate and lead citrate and viewed on a JEOL 1200 EX II transmission electron microscope (JEOL USA Inc., Peabody, MA). Images at magnifications of 3,000X to 30,000X were taken with an AMT 8-megapixel digital camera (Advanced Microscopy Techniques, Woburn, MA).

Immunofluorescence analysis

Formalin-fixed and paraffin-embedded sections were deparaffinized in xylene, rehydrated in an ethanol gradient, and boiled in a citrate buffer (Vector Laboratories Inc.) for antigen retrieval. After blocking with 2.5% goat serum in PBS (Vector laboratories) for 1 hr at room temperature, sections were incubated overnight at 4°C with primary antibodies (Supplementary Table 2) diluted in 2.5% normal goat serum. After washing with PBS, sections were incubated with Alexa Fluor 488-conjugated secondary antibodies (Life Technologies, Carlsbad, CA, USA) for 1 hr at room temperature, washed, and mounted with ProLong Gold Antifade Mountant with DAPI (Thermo Fisher Scientific). All Immunofluorescence images were obtained using a Zeiss LSM 880 confocal microscope (10x and 40x objective lens).

Fresh oviduct tissues from 8-week-old mice from control (n=3) and cKO (n=3) were collected and fixed in 4 % PFA at 4°C overnight. After that, oviducts were washed with PBS and cryoprotected using gradually increased sucrose concentrations in a row (15 and 30% w/v). 24 h later, oviducts were embedded in an OCT medium, frozen (−80 °C), and cryo-sectioned in the CM3050S Cryostat (Leica Biosystems, Germany). The cryosections were mounted on Superfrost plus glass slides (Fisher Scientific, Pittsburgh, PA) and stored at −80°C until used.

Cryo-sections were processed for immunofluorescence staining as described before [32]. Briefly, cryosections were blocked in 2.5% normal goat serum. After washing with PBS, sections were blocked with Mouse on Mouse (M.O.M.) IgG blocking reagent (M.O.M. Fluorescein Kit, Vector Laboratories, #FMK-2201) diluted in 1% BSA per the manufacturer’s instructions. After blocking, sections were incubated overnight at 4°C with primary antibodies (Supplementary Table 2) diluted in 2.5% normal goat serum. Next day, following washing with PBS, sections were incubated with Alexa Fluor 488-conjugated secondary antibodies (Life Technologies, Carlsbad, CA, USA) for 1 hr at room temperature, washed, and mounted with ProLong Gold Antifade Mountant with DAPI (Thermo Fisher Scientific). All Immunofluorescence images were imaged using a Nikon Fluorescent microscope.

Western blotting

Protein lysates (40 µg per lane) from uteri or oviducts were loaded on a 4-15% SDS-PAGE gel (Bio-Rad, Hercules, CA, USA), separated in 1X Tris-Glycine Buffer (Bio-Rad), and transferred to PVDF membranes via a wet electro-blotting system (Bio-Rad), all according to the manufacturer’s directions [33]. PVDF membranes were blocked for 1 hour in 5% non-fat milk in Tris-buffered saline containing 0.1% Tween-20 (TBS-T, Bio-Rad), then incubated overnight at 4 °C with antibodies listed in Supplementary Table 2 in 5% BSA in TBS-T. Blots were then probed with anti-Rabbit IgG conjugated with horseradish peroxidase (1:5000, Cell Signaling Technology, Danvers, MA, USA) in 5% BSA in TBS-T for 1 hr at room temperature. Signal was detected with the Pierce™ ECL Western Blotting Substrate (Millipore, Billerica, MA, USA), and blot images were collected with a Bio-Rad ChemiDoc imaging system.

RNA Isolation and Quantitative Real-Time RT-PCR Analysis

Tissues/cells were lysed in RNA lysis buffer, and total RNA was extracted with the Purelink RNA mini kit (Invitrogen, Carlsbad, CA, USA) according to the manufacturer’s instructions. RNA was quantified with a Nano-Drop 2000 (Thermo Fisher Scientific). Then, 1 µg of RNA was reverse transcribed with the High-Capacity cDNA Reverse Transcription Kit (Thermo Fisher Scientific). The amplified cDNA was diluted to 10 ng/µL, and qRT-PCR was performed with primers listed in the Supplementary Table 1 and TaqMan 2X master mix (Applied Biosystems/Life Technologies, Grand Island, NY, USA) on a 7500 Fast Real-time PCR system (Applied Biosystems/Life Technologies). The delta-delta cycle threshold method was used to normalize expression to the reference gene 18S.

Treatment of mice with Polyphyllin VI

Polyphyllin VI (Selleck chemicals, Houston, TX, USA), a pharmacological agent that induces caspase-1-mediated pyroptosis, was used to study its effects on embryo transport in the oviduct [34]. Eight-week-old C57BL/6 mice were injected for three consecutive days starting from 1 dpc with Polyphyllin VI activator dissolved in 40% PEG-300 (15mg/kg body weight). Dimethyl sulfoxide with 40% PEG-300 was administered as a vehicle.

Statistics

A two-tailed paired student t-test was used to analyze data from experiments with two experimental groups and one-way ANOVA followed by Tukey’s post hoc multiple range test was used for multiple comparisons. All data are presented as mean ±SEM. GraphPad Prism 9 software was used for all statistical analyses. Statistical tests, including p values, are reported in the corresponding figure legends or, when possible, directly on the data image. To ensure the reproducibility of our findings, experiments were replicated in a minimum of three independent samples, to demonstrate biological significance, and at least three independent times to ensure technical and experimental rigor and reproducibility.

Results

Conditional deletion of Atg14 in the FRT results in infertility despite the normal ovarian function

Earlier version: To explore the role of Atg14 in uterine function, we first determined its expression levels in uterus during early pregnancy (Day 1-7) in mice. We found distinct expression of ATG14 in all the uterine compartments (luminal epithelium, glands, and stroma) by day 1 of pregnancy, which disappeared by day 2 of pregnancy (Fig. 1A). However, the ATG14 expression in uterus reappeared by day 3 and persisted through the day 7 of pregnancy. The period from day 3 to day 7 is critical, as it is during this time when uterus begins to prepare for embryo implantation and undergoes decidualization process. Consistent with protein expression, similar Atg14 expression at mRNA levels was noted in uteri from early pregnant mice (Fig. 1B). This analysis suggests a potential role for ATG14 protein in the uterine physiologic adaptations during early pregnancy. Thus, to study the role of ATG14 in uterine function, we generated a conditional knockout (PRcre/+ /Atg14flox/flox) mouse model by crossing Atg14 flox/flox mice with mice expressing Cre recombinase under the control of progesterone receptor promoter (PRcre/+). Histological examination of the uterus from adult females showed no gross morphological differences between Atg14 cKO and control mice (Fig. S1A). Further, we did not find any overt defects in ovary as cKO mice had normal follicles and corpus luteum as like their corresponding controls (Fig. S1B). Analysis of transcript levels from the uterus, ovary, and liver showed that while Atg14 levels were efficiently depleted in uteri from cKO mice the levels were unaltered in the ovary and liver samples in control and cKO mice (Fig. 1C). Immunofluorescence analysis further confirmed the efficient deletion of ATG14 in all uterine compartments (epithelium and stroma) of cKO mice uteri compared to corresponding control groups (Fig. 1D).

Revised version: To explore the role of Atg14 in uterine function, we first determined its expression levels in the uterus during early pregnancy (Day 1-7) in mice. We found distinct expression of ATG14 in all the uterine compartments (luminal epithelium, glands, and stroma) by day 1 of pregnancy, which disappeared by day 2 of pregnancy (Fig. 1A). However, the ATG14 expression in the uterus reappeared by day 3 and persisted through the day 7 of pregnancy. The period from day 3 to day 7 is critical, as it is during this time when the uterus begins to prepare for embryo implantation and undergoes the decidualization process. Consistent with protein expression, similar Atg14 expression at mRNA levels was noted in uteri from early pregnant mice (Fig. 1B). This analysis suggests a potential role for ATG14 protein in the uterine physiologic adaptations during early pregnancy. Thus, to study the role of ATG14 in uterine function, we generated a conditional knockout (PRcre/+ /Atg14flox/flox) mouse model by crossing Atg14 flox/flox mice with mice expressing Cre recombinase under the control of progesterone receptor promoter (PRcre/+). Histological examination of the uterus from adult females showed no gross morphological differences between Atg14 cKO and control mice (Fig. S1A). Further, we did not find any overt defects in the ovary as cKO mice had normal follicles and corpus luteum like their corresponding controls (Fig. S1B). Analysis of transcript levels from the uterus, ovary, and liver showed that while Atg14 levels were efficiently depleted in uteri from cKO mice the levels were unaltered in the ovary and liver samples in control and cKO mice (Fig. 1C). Immunofluorescence analysis further confirmed the efficient deletion of ATG14 in all uterine compartments (epithelium and stroma) of cKO mice uteri compared to corresponding control groups (Fig. 1D).

Atg14 loss in the FRT results in infertility.

(A) Representative immunofluorescent images of uteri from pregnant mice (n=3) at the indicated days of pregnancy stained with an ATG14-specific antibody (green). LE: luminal epithelium, G: glands, S: stroma, Scale bar: 100 µm. Rabbit IgG was used as an isotype control for staining. (B) Relative transcript levels of Atg14 mRNA in uteri from pregnant mice (n=3-5) at indicated days of pregnancy. mRNA levels are normalized to levels of 18S m-RNA. Data are presented as mean ±SEM; **P<0.01, P>0.05, ns=non-significant. (C) Relative mRNA levels of Atg14 in 8-week-old virgin control and cKO mice uteri, ovary, and liver (n=5). mRNA levels are normalized to levels of 18S mRNA. Data are presented as mean ±SEM; ***P<0.001, P>0.05, ns=non-significant. (D) Representative immunofluorescent images of ATG14 expression in different uterine compartments in control (n=5) and Atg14 cKO mice (n=5). LE: luminal epithelium, G: glands, S: stroma. (E) (left panel) Relative number of pups/female/litter and (right panel) and relative number of total pups/months of Atg14 control (n=4) and cKO mice (n=5) sacrificed after the breeding trial. Data are presented as mean ±SEM; ***P<0.001; P>0.05, ns=non-significant.

Considering the effective ablation, a 6-month breeding study was performed mating virile wild-type male mice with adult Atg14 cKO and control female mice. We found that Atg14 cKO females did not deliver any litter during the 6-months trial period. However, the control female (Atg14 flox/flox) mice delivered an average of ∼7-8 pups per litter every month (Table 1) (Fig. 1E). These findings suggest an indispensable role of Atg14 in female fertility with intact ovarian function.

Loss of Atg14 results in impaired embryo implantation and uterine receptivity in mice

Based on the normal ovarian morphology, we posited that the infertility observed in females with Atg14 cKO status could be attributed to compromised uterine functioning. Analysis of the implanting embryos at Day 5 of pregnancy showed no implantation sites in the uteri of Atg14 cKO females, whereas ∼8 to 9 implantation sites were seen in the uteri of control mice at 5 dpc (Fig. 2A, left panel). Whilst there was no embryo present in the uterine lumen of Atg14 cKO mice, a fully attached embryo encapsulated by the luminal uterine epithelium, was seen in control mice uteri (Fig. 2A, middle panel). Additionally, MUC1, a receptive marker expression persisted in the luminal epithelium of cKO mice uteri at 5 dpc as well as 4 dpc (Fig. 2A, right panel, Fig. 2C). The process of successful embryo attachment and implantation within the uterus necessitates a transition from a non-receptive to a receptive state, a transformation orchestrated under the regulated influence of steroid hormones. Thus, the effects of Atg14 loss on uterine responsiveness to steroid hormones E2 and P4 were characterized. We performed an established hormone-induced uterine receptivity experiment, which involved ordered and co-stimulatory actions of E2 and P4 leading to the initiation of a receptive phase [35]. In response to E2 treatment, uteri from both Atg14 control and cKO mice showed similar proliferation as evidenced by Ki-67 staining (Fig. 2B, middle panel). Following the co-stimulation with E2+P4 treatment, epithelial proliferation was inhibited with concomitant induction of stromal cell proliferation in control mice uteri indicating fully receptive uteri (Fig. 2B). Interestingly, cKO mice uteri failed to elicit sub-epithelial stromal cell proliferation and showed intact P4-driven inhibition of epithelial proliferation. Consistently, uteri from Atg14 cKO mice at 4 dpc showed a reduced number of proliferating stromal cells (Fig. 2D). The absence of significant changes in E2-induced targets, such as Lif, Mcm2, Ccnd1, and Fgf18, at 4 dpc (Fig. 2E) supports our conclusion that ATG14 is required for P4-mediated but not for E2-mediated actions during uterine receptivity. Moreover, the normal serum levels of E2 and P4 at D4 of pregnancy rule out any hormonal imbalances, strongly suggesting that the observed phenotype is primarily due to the uterine-specific loss of Atg14 (Fig. 2F).

Atg14 is critical for embryo implantation, and uterine receptivity.

Earlier version: (A) Gross images of 5.0 dpc uteri of control (n=5) and Atg14 cKO mice (n=5) injected with Chicago Sky Blue dye to visualize implantation sites (denoted by black arrows) (left panel). H&E-stained cross-sections (4X & 40X) of 5.0 dpc uteri of control (n=5) and Atg14 cKO (n=5) mice to visualize embryo implantation (Middle panel). The asterisk denotes the embryo. Immunofluorescence analysis of uterine tissues from control (n=5) and Atg14 cKO mice (n=5), stained with MUC1 and KRT8 (right panel). LE: luminal epithelium, G: glands, S: stroma. (B) Representative immunofluorescence images of uteri from control (n=5) and Atg14 cKO mice (n=5) stained for Ki-67 following Oil or E2 or E2+P4 treatment (n=5 mice/group); scale bar: 100 μm. LE: luminal epithelium, G: glands, S: stroma. (C) Immunofluorescence analysis of KRT8 (green), MUC1 (Red), and (D) Ki-67 (red) in the uteri of 4 dpc control and Atg14 cKO mice; scale bars, 100 µm. (E) Relative transcript levels of Lif, Mcm2, Ccnd1, and Fgf18 in control and cKO uteri at 4dpc. mRNA levels are normalized to levels of 18S m-RNA. Data are presented as mean ±SEM; ns=non-significant. (F) Levels of steroid hormones estradiol and progesterone from serum collected during euthanasia of 4 dpc control or cKO mice.

Revised version: (A) Gross images of 5.0 dpc uteri of control (n=5) and Atg14 cKO mice (n=5) injected with Chicago Sky Blue dye to visualize implantation sites (denoted by black arrows) (left panel). H&E-stained cross-sections (6X & 20X) of 5.0 dpc uteri of control (n=5) and Atg14 cKO (n=5) mice to visualize embryo implantation (Middle panel). The asterisk denotes the embryo. Immunofluorescence analysis of uterine tissues from control (n=5) and Atg14 cKO mice (n=5), stained with MUC1 and KRT8 (right panel). LE: luminal epithelium, G: glands, S: stroma. (B) Representative immunofluorescence images of uteri from control (n=5) and Atg14 cKO mice (n=5) stained for Ki-67 following Oil or E2 or E2+P4 treatment (n=5 mice/group); scale bar: 100 μm. LE: luminal epithelium, G: glands, S: stroma. (C) Immunofluorescence analysis of KRT8 (green), MUC1 (Red), and (D) Ki-67 (red) in the uteri of 4 dpc control and Atg14 cKO mice; scale bars, 100 µm. (E) Relative transcript levels of Lif, Mcm2, Ccnd1, and Fgf18 in control and cKO uteri at 4dpc. mRNA levels are normalized to levels of 18S m-RNA. Data are presented as mean ±SEM; ns=non-significant. (F) Levels of steroid hormones estradiol and progesterone from serum collected during euthanasia of 4 dpc control or cKO mice.

Atg14 is required for maintaining oviductal cell structural integrity and embryo transport

Earlier version: Given that Atg14 cKO mice had impaired embryo implantation, we wondered whether embryos are reaching the uterus timely through the oviduct. To determine this, we flushed embryos from both the control and cKO mice uteri on day 4 of pregnancy. Interestingly, in cKO mice, we could retrieve only 1-2% of embryos from their uteri, whereas in control mice, 100% of well-developed blastocysts were retrieved from their uteri (Fig. 3A & B). To ensure the timely transport of all embryos from the oviducts to the uteri, we also flushed oviducts from both control and cKO mice. As expected, in the control mice oviduct flushing, we could not recover any embryos or blastocysts, indicating their precise and timely transport to the uterus. Unexpectedly, oviduct flushing from cKO mice resulted in the retrieval of approximately 80-90% of blastocysts, suggesting their potential entrapment within the oviducts, impeding their transit to the uterus (Fig. 3A & B). At 4 dpc, there was no significant difference in the average number of blastocysts, morula, or non-viable or retrieved from Atg14 control uteri and cKO oviducts (Fig. S4A). However, we noted that the percentage of developmentally delayed embryos appeared to be higher in Atg14 cKO oviducts compared to the embryos retrieved from control uteri (Fig. 3C). The histological analysis further confirmed an entrapped embryo in the ampulla of Atg14 cKO at 4 dpc as shown in Fig. 3D. Given the embryo retention phenotype in oviduct, we sought to determine if ATG14 is expressed in this region. Consistent with previous studies reporting PR-cre activity in the isthmus [8, 29], we found a significant depletion of ATG14 in the isthmus compared to the ampulla (Fig. S2A). Further, the increased p62 and LC3B expression in cKO oviducts suggests that the observed embryo retention phenotype might be attributed to the loss of ATG14-dependent autophagy in cKO oviducts (Fig. S2B).

Revised version: Given that Atg14 cKO mice had impaired embryo implantation, we wondered whether embryos are reaching the uterus timely through the oviduct. To determine this, we flushed embryos from both the control and cKO mice uteri on day 4 of pregnancy. Interestingly, in cKO mice, we could retrieve only 1-2% of embryos from their uteri, whereas in control mice, 100% of well-developed blastocysts were retrieved from their uteri (Fig. 3A & B). To ensure the timely transport of all embryos from the oviducts to the uteri, we also flushed oviducts from both control and cKO mice. As expected, in the control mice oviduct flushing, we could not recover any embryos or blastocysts, indicating their precise and timely transport to the uterus. Unexpectedly, oviduct flushing from cKO mice resulted in the retrieval of approximately 90% of embryos, suggesting their potential entrapment within the oviducts, impeding their transit to the uterus (Fig. 3A & B). At 4 dpc, there was no significant difference in the average number of blastocysts, morula, or non-viable or retrieved from Atg14 control uteri and cKO oviducts (Fig. S4A). However, we noted that the percentage of developmentally delayed embryos appeared to be higher in Atg14 cKO oviducts compared to the embryos retrieved from control uteri (Fig. 3C). The histological analysis further confirmed an entrapped embryo in the ampulla of Atg14 cKO at 4 dpc as shown in Fig. 3D. Given the embryo retention phenotype in oviduct, we sought to determine if ATG14 is expressed in this region. Consistent with previous studies reporting PR-cre activity in the isthmus [8, 29], we found a significant depletion of ATG14 in the isthmus compared to the ampulla (Fig. S2A). Further, the increased p62 and LC3B expression in cKO oviducts suggests that the observed embryo retention phenotype might be attributed to the loss of ATG14-dependent autophagy in cKO oviducts (Fig. S2B).

Atg14 is critical for embryo transport in oviduct.

(A) Representative images (upper panel) and (B) Percentage of embryos (lower panel) collected at 4 dpc from the uteri or the oviducts of control (n=6) or Atg14 cKO mice (n=6-8). (C) Percentage of blastocysts, morulae, developmentally delayed or nonviable embryos collected from Atg14 control mice uteri and Atg14 cKO mice oviducts at 4 dpc. (D) Histological analysis using H&E staining of the ampullary and isthmic region of the oviduct from control and Atg14 cKO female mice at 4 dpc (n=3 mice/genotype). (E) Embryos retrieved from the oviduct and uterus of super-ovulated Atg14 control or cKO mice at 4 dpc (n=3 mice/genotype). (F) Immunofluorescence analysis of KRT8 (green), and α-SMA (Red), in the oviduct of 4 dpc control (n=5) and Atg14 cKO mice (n=5) ampulla (upper panel) and isthmus (lower panel).

Similarly, when we super-ovulated the mice and looked for the embryos on day 4 of pregnancy, we could recover ∼20-25 embryos from the cKO mice oviducts compared to only 4-5 embryos that were able to reach the uterus. In comparison, in control mice, 90% of embryos were able to complete their pre-destined and timely transport to the uterus (Fig. 3E), except 1-2% of unfertilized embryos, which remained in their oviducts.

To understand the underlying cause for retained embryos in cKO mice oviducts, we performed histology analysis and determined the structural morphology of their oviducts. The cKO mice oviduct lining shows marked eosinophilic cytoplasmic change akin to decidualization in human oviducts. Some of the cells showed degenerative changes with cytoplasmic vacuolization and nuclear pyknosis, loss of nuclear polarity, and loss of distinct cell borders giving an appearance of fusion of cells (Fig. 3F). The marked cytologic enlargement appears to cause luminal obliteration or narrowing resulting in a completely unorganized, obstructed, and narrow lumen, thereby, hampering the path of embryos to reach the uterus as evident from an entrapped embryo shown in Fig. 3D. Further, epithelial (cytokeratin KRT8-positive) and myosalpinx (α-smooth muscle active [SMA]-positive,) marker analysis revealed completely distorted epithelial structures (in terms of loss of epithelial cell integrity) with no overt defects in the muscle organization in oviducts of cKO mice (Fig. 3F, lower panel). Notably, the ampullary region from cKO oviducts seems to be normal with intact epithelial and smooth muscle structures as evident from KRT8 and SMA staining (Fig. 3F, upper panel).

Selective loss of Atg14 in oviduct cilia is dispensable for female fertility

Earlier version: The oviduct epithelium primarily consists of two types of cells: ciliated and secretory (non-ciliated) cells. Ciliated cells play a role in embryo and oocyte transport by means of ciliary beat, and non-ciliated/secretory cells produce an oviductal fluid that is rich in amino acids and various molecules, thereby providing an optimal micro-environment for sperm capacitation, fertilization, embryonic survival, and development [36]. Therefore, we determined whether oviducts from cKO mice possess normal ciliated and secretory cell composition. To do so, we examined the expression of various markers, including acetylated-α-tubulin (cilia marker), FOXJ1 (ciliogenesis markers), and PAX8 (non-ciliated/ secretory cell marker) in control and cKO oviducts. As shown in Fig. 4A, while we observed normal ciliary structures in the ampulla of both control and cKO oviducts, there was a substantial loss of the ciliary epithelial cells (indicated by fewer α-tubulin and FOXJ1-positive cells) (Fig. 4B, left panel and Fig. S3) as well as secretory cells (indicated by fewer PAX-8 positive cells) in the isthmus of cKO oviducts (Fig. 4B, right panel).

Revised version: The oviduct epithelium primarily consists of two types of cells: ciliated and secretory (non-ciliated) cells. Ciliated cells play a role in embryo and oocyte transport by means of ciliary beat, and non-ciliated/secretory cells produce an oviductal fluid that is rich in amino acids and various molecules, thereby providing an optimal micro-environment for sperm capacitation, fertilization, embryonic survival, and development [36]. Therefore, we determined whether oviducts from cKO mice possess normal ciliated and secretory cell composition. To do so, we examined the expression of various markers, including acetylated-α-tubulin (cilia marker), FOXJ1 (ciliogenesis markers), and PAX8 (non-ciliated/ secretory cell marker) in control and cKO oviducts. As shown in Fig. 4A, normal ciliary structures were observed in the ampulla of both control and cKO oviducts. However, in the isthmus of cKO oviducts, we observed a reduction in both FOXJ1- and PAX8-expressing cells (Fig. 4B, and Fig. S3). Given the importance of cilia in embryo transport, we wondered whether the loss of Atg14 in oviduct cilia has any impact on embryo transport. To address this, we generated a Foxj1-cre/+; Atg14 f/f mouse model, wherein Atg14 will be ablated only in ciliary epithelial cells. We observed that ciliated epithelial cells that were positive for acetylated α-tubulin staining did not appear to be different in Foxj1Cre/+; Atg14 f/f mice oviduct compared to controls suggesting normal ciliogenesis in Foxj1Cre/+; Atg14 f/f mice (Fig. 4C).

Atg14 loss in oviduct cilia is dispensable for embryo transport.

(A) Immunofluorescence analysis of acetylated α-tubulin (green), and DAPI (blue) in oviduct of 4 dpc control (n=5) and Atg14 cKO mice (n=5). (B) Immunohistochemical analysis of FOXJ1 and PAX8 at 4 dpc (n=5). (C) Immunohistochemical analysis of acetylated α-tubulin in 8-week-old control (n=5) and Foxj1/Atg14 cKO mice (n=5). Images are taken at 6X and 40X. Scale bar: 100µm and 10 µm. (D & E) Relative number of pups/females/litter and relative number of total pups/months of control (n=5) and Foxj1/Atg14 cKO mice (n=5) sacrificed after the breeding trial. Data are presented as mean ± SEM; P>0.05, ns=non-significant. (F) Gross images of 5.0 dpc uteri of control (n=3) and Foxj1/Atg14 cKO mice (n=3) injected with Chicago Sky Blue dye to visualize implantation sites (denoted by black arrows).

Removed in the revised version: Given the importance of cilia in embryo transport, we wondered whether the loss of Atg14 in oviduct cilia has any impact on embryo transport. To address this, we generated a Foxj1-cre/+; Atg14 f/f mouse model, wherein Atg14 will be ablated only in ciliary epithelial cells. We observed that ciliated epithelial cells that were positive for acetylated α-tubulin staining did not appear to be different in Foxj1Cre/+; Atg14 f/f mice oviduct compared to controls suggesting normal ciliogenesis in Foxj1Cre/+; Atg14 f/f mice (Fig. 4C).

The 6-month breeding trial revealed that loss of Atg14 in oviductal cilia had no impact on fertility (Table 2, Fig. 4D & E). Consistently, D5 implantation study analysis showed ∼7-8 visible embryo implantation sites in Foxj1/Atg14 cKO mice like their corresponding controls suggesting that embryos were able to make their way to the uterus in a timely manner and undergo implantation despite the ablation of Atg14 in the oviduct ciliary epithelial cells (Fig. 4F). These findings suggest that ciliary expression of Atg14 is dispensable for embryo transport.

ATG14 maintains mitochondria integrity and prevents unscheduled pyroptosis activation in the oviduct to enable embryo transport

Earlier version: To gain more insights into structural defects, we performed TEM analysis on oviducts collected from both control and cKO females on day 4 of pregnancy. Interestingly, oviducts from cKO females had numerous altered mitochondrial structures with abnormally enlarged mitochondria and a less dense matrix compared to control oviducts that had small and compact mitochondria with tight cristae and a dense matrix (Fig. 5A). Corroborating and extending those findings, we observed marked reductions of mitochondria network (as defined by TOMM20 staining) in the sub-nuclear region of oviducts from Atg14 cKO mice compared to more densely packed mitochondrial network in the sub-nuclear region of control ones (Fig. 5B). Further co-localization study analysis revealed a uniform co-localization of Cytochrome C with the TOM20 positive mitochondrial network in control oviducts reflecting the intact mitochondrial integrity (Fig. 5B). In contrast, in Atg14 cKO oviducts, cells with disrupted mitochondrial networks exhibited increased cytosolic leakage of Cytochrome C (Fig. 5B). Additionally, analysis of various mitochondrial functional and architecture markers (Cox4i2, Pink1, Opa1, and Drp1 showed reduced expression in Atg14 cKO compared to controls (Fig. 5C).

Revised version: To gain more insights into structural defects, we performed TEM analysis on oviducts collected from both control and cKO females on day 4 of pregnancy. Interestingly, oviducts from cKO females had numerous altered mitochondrial structures with abnormally enlarged mitochondria and a less dense matrix compared to control oviducts that had small and compact mitochondria with tight cristae and a dense matrix (Fig. 5A). Corroborating and extending those findings, we observed marked reductions of mitochondria network (as defined by TOM20 staining) in the sub-nuclear region of oviducts from Atg14 cKO mice compared to more densely packed mitochondrial network in the sub-nuclear region of control ones (Fig. 5B). Further co-localization study analysis revealed a uniform co-localization of Cytochrome C with the TOM20 positive mitochondrial network in control oviducts reflecting the intact mitochondrial integrity (Fig. 5B). In contrast, in Atg14 cKO oviducts, cells with disrupted mitochondrial networks exhibited increased cytosolic leakage of Cytochrome C (Fig. 5B). Additionally, analysis of various mitochondrial functional and architecture markers (Cox4i2, Pink1, Opa1, and Drp1 showed reduced expression in Atg14 cKO compared to controls (Fig. 5C).

Atg14 facilitates embryo transport in the oviduct by preserving mitochondrial integrity and inhibiting the activation of pyroptosis.

Earlier version: (A) Transmission electron microscopy of oviducts at 4 dpc from control (n=3) and Atg14 cKO (n=3). Black arrowheads in control oviducts show small, compact mitochondria with tight cristae. Red arrowheads in cKO oviducts show abnormally enlarged mitochondria with loose cristae. (B) Immunofluorescence analysis of TOM20 (green) and Cytochrome C (red) in Atg14 control and cKO oviducts cryo-sections. The inset shows the zoomed image of the selected area from control and cKO oviducts, Arrowheads in cKO inset oviduct section show the leaked cytochrome C in the cytoplasm. Tissues were counterstained with DAPI (blue) to visualize nuclei; scale bars, 10 µm. (C) Relative transcript levels of Opa1, Cox4i2, Drp1, and Pink1 in oviduct tissues. Data are presented as mean ±SEM. *P<0.05; **P<0.01 compared with controls. 18S was used as an internal control (D) Immunofluorescence analysis of GSDMD (red) + KRT8 (green), Caspase-1 (red) + KRT8 (green) in oviducts of adult control (n=5), and Atg14 cKO mice (n=5). Tissues were counterstained with DAPI (blue) to visualize nuclei; scale bars, 100 µm. (E) Immunohistochemical analysis of GSDMD expression in adult oviduct tissues (left panel). The middle and -right panels show the zoom-in images to show the relative GSDMD expression in the ampulla and isthmus section from control and cKO oviducts. (F) Immunohistochemical analysis of GSDMD expression in adult uterus and ovary tissues (G) Western blotting to show protein levels of Caspase-1, and GSDMD in oviduct tissues. β-actin is used as a loading control. (H) Relative transcript levels of Tnf-α and Cxcr3 in oviduct tissues. Data are presented as mean ±SEM. *P<0.05; **P<0.01; ***P<0.001 compared with controls. 18S was used as an internal control.

Revised version: (A) Transmission electron microscopy of oviducts at 4 dpc from control (n=3) and Atg14 cKO (n=3). Black arrowheads in control oviducts show small, compact mitochondria with tight cristae. Red arrowheads in cKO oviducts show abnormally enlarged mitochondria with loose cristae. (B) Immunofluorescence analysis of TOM20 (green) and Cytochrome C (red) in Atg14 control and cKO oviducts cryo-sections. The inset shows the zoomed image of the selected area from control and cKO oviducts, Arrowheads in the cKO inset oviduct section show the leaked cytochrome C in the cytoplasm. Tissues were counterstained with DAPI (blue) to visualize nuclei; scale bars, 10 µm. (C) Relative transcript levels of Opa1, Cox4i2, Drp1, and Pink1 in oviduct tissues. Data are presented as mean ±SEM. *P<0.05; **P<0.01 compared with controls. 18S was used as an internal control (D) Immunofluorescence analysis of GSDMD (red) + KRT8 (green), Caspase-1 (red) + KRT8 (green) in oviducts of adult control (n=5), and Atg14 cKO mice (n=5). Tissues were counterstained with DAPI (blue) to visualize nuclei; scale bars, 100 µm. (E) Immunohistochemical analysis of GSDMD expression in adult oviduct tissues (left panel). The middle and -right panels show the zoom-in images to show the relative GSDMD expression in the ampulla and isthmus section from control and cKO oviducts. (F) Immunohistochemical analysis of GSDMD expression in adult uterus and ovary tissues (G) Western blotting to show protein levels of Caspase-1, and GSDMD in oviduct tissues. β-actin is used as a loading control. (H) Relative transcript levels of Tnf-α and Cxcr3 in oviduct tissues. Data are presented as mean ±SEM. *P<0.05; **P<0.01; ***P<0.001 compared with controls. 18S was used as an internal control.

Earlier version: Having identified the cellular mechanisms of Atg14, we aimed to delineate the underlying molecular mechanisms associated with Atg14 function in the oviduct. Since we found mitochondrial defects with ATG14 loss and altered mitochondrial structures have been linked to the activation of pyroptosis, we determined if ATG14 regulates pyroptosis in oviduct [37–39]. First, we assessed the key primary executors of the pyroptosis pathway, Gasdermin D (GSDMD), and caspase-1. Immunofluorescence analysis revealed a remarkable upregulation in GSDMD and caspase-1 expression in the cKO oviducts compared to controls (Fig. 5D & E). However, histological analysis of uterus and ovary showed no induction in GSDMD expression compared to their corresponding controls (Fig. 5F). Western blot analysis further confirmed elevated expression of caspase-1 and GSDMD in cKO oviducts in comparison to control oviducts as shown in Fig. 5G. Additionally, the qPCR analysis demonstrated elevated levels of inflammatory markers, such as Tnf-α and Cxcr3 in cKO oviducts compared to control ones (Fig. 5H). Based on these findings, we posit that Atg14 plays a crucial role in regulating the pyroptotic pathway by preserving the mitochondrial structural and functional integrity, and activation of pyroptosis owing to loss of ATG14 within the oviduct.

Revised version: Having identified the cellular mechanisms of Atg14, we aimed to delineate the underlying molecular mechanisms associated with Atg14 function in the oviduct. Since we found mitochondrial defects with ATG14 loss and altered mitochondrial structures have been linked to the activation of pyroptosis, we determined if ATG14 regulates pyroptosis in the oviduct [37–39]. First, we assessed the key primary executors of the pyroptosis pathway, Gasdermin D (GSDMD), and caspase-1. Immunofluorescence analysis revealed a remarkable upregulation in GSDMD and caspase-1 expression in the cKO oviducts compared to controls (Fig. 5D & E). However, histological analysis of uterus and ovary showed no induction in GSDMD expression compared to their corresponding controls (Fig. 5F). Western blot analysis further confirmed elevated expression of caspase-1 and GSDMD in cKO oviducts in comparison to control oviducts as shown in Fig. 5G. Additionally, the qPCR analysis demonstrated elevated levels of inflammatory markers, such as Tnf-α and Cxcr3 in cKO oviducts compared to control ones (Fig. 5H). Based on these findings, we posit that Atg14 plays a crucial role in regulating the pyroptotic pathway by preserving the mitochondrial structural and functional integrity in the oviduct.

Earlier version: To substantiate the notion, we also evaluated the impact of unscheduled activation of pyroptosis on embryo transport. To test this, we employed a pyroptosis inducer, Polyphyllin VI, and chose the optimal dose based on the established studies [34]. Wild-type females were mated with virile males and following the plug detection treated with Polyphyllin VI for three consecutive days from 1 to 4 dpc (Fig. 5E). We chose to treat to 1 to 4 dpc for treatment to activate unwarranted pyroptosis in oviduct during the embryo transport. Following the treatments, the oviducts, and uteri from both vehicle-treated and Polyphyllin VI-treated were flushed. We found that pregnant females treated with Polyphyllin VI showed ∼50% embryo retention in the oviduct, whereas in the vehicle-treated group, no embryos were retained in the oviduct (Fig. 6A-C). Further analysis of embryos retrieved from Polyphyllin-treated oviducts showed more percentage of developmentally delayed and non-viable embryos compared to embryos recovered from control uteri (Fig. 6D). The average number of embryos recovered from polyphyllin- and vehicle-treated mice was not significantly different (Fig. S4B). Histological analysis showed a marked induction in GSDMD expression compared to vehicle-treated mice oviducts (Fig. 6E). However, precisely activating the unscheduled pyroptosis during the critical period of embryo transport is technically challenging. Nonetheless, our findings provide evidence that unscheduled pyroptosis adversely affects embryo transport through the oviduct. Taken together, these results demonstrate that Atg14 safeguards pyroptosis activation in the oviduct and allows the smooth transport of embryos to the uterus (Fig. 6F).

Revised version: To substantiate the notion, we also evaluated the impact of unscheduled activation of pyroptosis on embryo transport. To test this, we employed a pyroptosis inducer, Polyphyllin VI, and chose the optimal dose based on the established studies [34]. Wild-type females were mated with virile males and following the plug detection treated with Polyphyllin VI for three consecutive days from 1 to 4 dpc (Fig. 6A). We chose to treat mice from 1 to 4 dpc to activate unwarranted pyroptosis in oviduct during the embryo transport. Following the treatments, the oviducts, and uteri from both vehicle-treated and Polyphyllin VI-treated were flushed. We found that pregnant females treated with Polyphyllin VI showed ∼50% embryo retention in the oviduct, whereas in the vehicle-treated group, no embryos were retained in the oviduct (Fig. 6A-C). Further analysis of embryos retrieved from Polyphyllin-treated oviducts showed more percentage of developmentally delayed and non-viable embryos compared to embryos recovered from control uteri (Fig. 6D). The average number of embryos recovered from polyphyllin- and vehicle-treated mice was not significantly different (Fig. S4B). Histological analysis showed a marked induction in GSDMD expression compared to vehicle-treated mice oviducts (Fig. 6E). However, precisely activating the unscheduled pyroptosis during the critical period of embryo transport is technically challenging. Nonetheless, our findings provide evidence that unscheduled pyroptosis adversely affects embryo transport through the oviduct. Taken together, these results demonstrate that Atg14 safeguards pyroptosis activation in the oviduct and allows the smooth transport of embryos to the uterus (Fig. 6F).

Pharmacological activation of pyroptosis in the oviduct inhibits embryo transport.

Earlier version: (A) Experimental strategy for pyroptosis activation in pregnant female mice (B) Embryos flushed from the vehicle (n=3) or polyphyllin VI treated (n=3) D-4 pregnant females. (C) Percentage of embryos recovered from oviducts or uteri. (D) Percentage of blastocysts, morulae, developmentally delayed or nonviable embryos collected from vehicle or Polyphyllin IV oviducts at 4 dpc. (E) Immunohistochemistry to show GSDMD expression in the isthmus section of polyphyllin IV-treated and vehicle-treated mice. 40X objective, Scale bar: 5 µm. (F) Graphical illustration to show embryo transport and pyroptosis regulation in the oviduct.

Revised version: (A) Experimental strategy for pyroptosis activation in pregnant female mice (B) Embryos flushed from the vehicle (n=3) or polyphyllin VI treated (n=3) D-4 pregnant females. (C) Percentage of embryos recovered from oviducts or uteri. (D) Percentage of blastocysts, morulae, developmentally delayed or nonviable embryos collected from vehicle uteri or Polyphyllin IV oviducts at 4 dpc. (E) Immunohistochemistry to show GSDMD expression in the isthmus section of polyphyllin IV-treated and vehicle-treated mice. 40X objective, Scale bar: 5 µm. (F) Graphical illustration to show embryo transport and pyroptosis regulation in the oviduct.

Discussion

In this study, we delineated the essential role of ATG14 in maintaining the structural integrity of the oviduct by preventing pyroptosis, which enables smooth embryo transit during early pregnancy. Specifically, we dissected the tissue-specific functions of ATG14 in the FRT using Cre driver mice targeting the female reproductive tract. Given that PR-Cre expresses postnatally in different tissues of FRT such as corpus luteum, oviducts, and different cellular compartments (epithelial, stromal, and myometrium) of uteri [29], we found that conditional ablation of ATG14-mediated functions in FRT resulted in infertility owing to hampered transport of embryos from the oviduct. Additionally, Atg14 loss in uteri causes impaired embryo implantation and receptivity. Its ablation in the oviduct might lead to the activation of pyroptosis, an inflammatory-associated apoptosis pathway that causes the retention of embryos in the oviduct and prevents their timely transport to the uterus. However, the ovarian-specific functions were found to be unaffected despite the loss of Atg14 in the corpus luteum.

The human endometrium is a complex dynamic tissue that undergoes sequential phases of proliferation and differentiation to support embryo implantation during the conceptive. However, in the absence of an implanting embryo, the endometrium sheds (menstruation) and initiates the regeneration process [40]. Previous reports indicated that autophagy is modulated in the human endometrium during the menstrual cycle. For example, autophagy is altered during the proliferative and secretory phases of the menstrual cycle with its highest activity occurring during the secretory phase when the stroma is decidualized [41]. Similarly, the level of autophagy was higher in postmenopausal human uterine epithelial cells compared to premenopausal uterine epithelial cells indicating the onset of autophagy upon estrogen deprivation [42]. Although all these studies reported that autophagy is hormonally regulated, however, the role of autophagy-specific proteins in the endometrium was not established till our group’s recent reports. Specifically, studies from our group reported the role of three different autophagy-associated proteins: FIP200 [25], ATG16L [27], and BECLIN-1 in uterine functions [26]. Conditional ablation of Fip200 and Atg16l in the uterus displayed fertility defects owing to impaired implantation, uterine receptivity, and decidualization defects. On the contrary, loss of Beclin1 in the uterus caused progressive loss of endometrial progenitor stem cells resulting in severe uterine developmental defects and rendering the mice infertile [26]. Although the autophagy-related proteins we studied so far influenced uterine functions, [25–27], we found a distinct role for ATG14 in both the uterine and oviduct-specific functions. It is intriguing to note that the absence of ATG14 did not affect the tissue integrity of the uterus. However, the severe structural abnormalities in the oviduct due to Atg14 ablation unearthed a unique function of Atg14 in maintaining oviduct homeostasis. This distinctive role of ATG14, unlike other autophagy proteins, might be due to its pivotal role in the assembly of PtdIns3K complexes which is not the case for either FIP200 or ATG16L [40, 43, 44]. Nevertheless, understanding the specific contributions of each core autophagy protein in reproductive tract functions is necessary which requires substantial efforts.

The retention of preimplantation embryos in the oviduct has been established as a significant contributor to implantation failures, presenting challenges to the overall reproductive health of women [8–10]. Embryo transport in the oviduct is known to be controlled primarily by two major physiological responses: ciliary activity and muscle contractility. In our study, a dramatic decrease in the number of Foxj1+ve ciliary epithelial cells and Pax-8+ve secretory cells in Atg14 cKO mice implied cell-type specific actions for ATG14 in the oviduct. Interestingly, the specific ablation of Atg14 in Foxj1+ve ciliary epithelial cells of the oviduct does not appear to impact fertility. This suggests that the absence of ATG14 within ciliary cells does not impact the process of embryo transport. Although unexpected, this is consistent with other reports that have demonstrated the dispensability of cilia for embryo transport in the oviduct [8, 9].

Earlier version: Pyroptosis is a highly inflammatory form of programmed cell death, characterized by cell swelling, flattening of the cytoplasm, and large bubbles-like protrusions on the plasma membrane [45–47]. Several recent studies found a link between autophagy and pyroptosis [48]. In vivo studies also found that mice lacking the autophagy genes Atg14, Fip200, Atg5, or Atg7 in myeloid cells had more pronounced lung inflammation [49]. In our study, cellular swelling, and fused membranous structures (a unique feature of activation of pyroptosis) observed in the oviductal epithelial folds from Atg14 cKO mice. This aberrant activation of the pyroptosis pathway due to loss of Atg14 appears to adversely affect the oviduct cellular structural integrity. These findings have implications beyond the oviduct, as autophagy can modulate the inflammatory signaling pathway through different avenues. For example, first, autophagy prevents mitochondrial reactive oxygen species release that can activate the inflammasome by inhibiting IL1-β and IL18 production through the digestion of dysfunctional mitochondria [50]. Second, autophagy is capable of targeting inflammasome complexes for degradation, which prevents the cleavage of pro-IL-1β and pro-IL-18 into biologically active forms. Finally, autophagy machinery further regulates IL-1β levels by engulfing and degrading pro-IL-1β proteins [50]. In our study, accumulation of abnormally enlarged and dysfunctional mitochondria in the absence of ATG14 in the oviduct clearly suggests the induction of severe inflammatory conditions hampering the embryo transit through the oviduct. These findings are further substantiated by the upregulation of pivotal mediators of the pyroptosis program, including GSDMD and caspase-1, in Atg14 cKO mice. Notably, the induction of pyroptosis in the absence of ATG14 is specifically localized to the oviducts, with no parallel activation detected in either the ovary or uterus. One potential reason for this observed phenomenon could be the oviduct’s unique molecular signatures that get activated in response to the presence of gametes. A recent study by Finnerty et al. (2024) reported that sperm induce pro-inflammatory conditions in the oviduct, which are preceded by an anti-inflammatory response triggered by the presence of embryos [51]. The persistent activation of inflammatory or pyroptotic conditions following the loss of Atg14 suggests that Atg14 may be a critical autophagy protein responsible for suppressing pyroptosis in the oviduct. Moreover, understanding the complex rheostat between the ATG14 and inflammation regulatory axis is highly intriguing and will necessitate future studies employing sophisticated models, such as combined knockout mice where ATG14 is deleted alongside key inflammatory regulators (e.g., NLRP3, GSDMD, or CASPASE-1).

Revised version: Pyroptosis is a highly inflammatory form of programmed cell death, characterized by cell swelling, flattening of the cytoplasm, and large bubbles-like protrusions on the plasma membrane [45–47]. Several recent studies found a link between autophagy and pyroptosis [48]. In vivo studies also found that mice lacking the autophagy genes Atg14, Fip200, Atg5, or Atg7 in myeloid cells had more pronounced lung inflammation [49]. In our study, cellular swelling, and fused membranous structures (a unique feature of activation of pyroptosis) were observed in the oviductal epithelial folds from Atg14 cKO mice. This aberrant activation of the pyroptosis pathway due to loss of Atg14 appears to adversely affect the oviduct cellular structural integrity. These findings have implications beyond the oviduct, as autophagy can modulate the inflammatory signaling pathway through different avenues. For example, first, autophagy prevents mitochondrial reactive oxygen species release that can activate the inflammasome by inhibiting IL1-β and IL18 production through the digestion of dysfunctional mitochondria [50]. Second, autophagy is capable of targeting inflammasome complexes for degradation, which prevents the cleavage of pro-IL-1β and pro-IL-18 into biologically active forms. Finally, autophagy machinery further regulates IL-1β levels by engulfing and degrading pro-IL-1β proteins [50]. In our study, accumulation of abnormally enlarged and dysfunctional mitochondria in the absence of ATG14 in the oviduct clearly suggests the induction of severe inflammatory conditions hampering the embryo transit through the oviduct. These findings are further substantiated by the upregulation of pivotal mediators of the pyroptosis program, including GSDMD and caspase-1, in Atg14 cKO mice. Notably, the induction of pyroptosis in the absence of ATG14 is specifically localized to the oviducts, with no parallel activation detected in either the ovary or uterus. One potential reason for this observed phenomenon could be the oviduct’s unique molecular signatures that get activated in response to the presence of gametes. A recent study by Finnerty et al. (2024) reported that sperm induce pro-inflammatory conditions in the oviduct, which are preceded by an anti-inflammatory response triggered by the presence of embryos [51]. The persistent activation of inflammatory or pyroptotic conditions following the loss of Atg14 suggests that Atg14 may be a critical autophagy protein responsible for suppressing pyroptosis in the oviduct. Moreover, understanding the complex rheostat between the ATG14 and inflammation regulatory axis is highly intriguing and will necessitate future studies employing sophisticated models, such as combined knockout mice where ATG14 is deleted alongside key inflammatory regulators (e.g., NLRP3, GSDMD, or CASPASE-1).

In this report, we revealed tissue-specific roles of ATG14 in governing oviductal transport, and uterine receptivity which are necessary for embryo implantation and pregnancy establishment. Of mechanistic insights, we delineated a novel function of ATG14 as a negative regulator of pyroptosis, which is vital for proper embryo transport in the oviduct (Fig. 6F). Thus, this study not only sheds light on the involvement of autophagy in oviductal embryo transport but also underscores the importance of autophagy in preserving the oviduct tissue integrity. Such insights hold promise for advancing our understanding of gynecological pathologies associated with the oviduct including tubal pregnancies.